<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week7/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

---
## Question 1:
#### Bootstrapping
Write a Python function to estimate the `mean` of the `heights` using `bootstrapping`. Your function should take the sample data as input and generate multiple bootstrap samples to compute the estimates.

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def bootstrap_estimate_height(sample_data, num_samples, sample_size):
    """
    Estimate the mean of heights using bootstrapping.

    Parameters:
    - sample_data (numpy.ndarray): Sample data containing heights.
    - num_samples (int): Number of bootstrap samples to generate.
    - sample_size (int): Size of each bootstrap sample.

    Returns:
    - mean_estimate (float): Estimated mean height.
    """
    # YOUR CODE HERE
    mean_estimates = []
    for _ in range(num_samples):
        bootstrap_sample = np.random.choice(sample_data, size=sample_size, replace=True)
        mean_estimates.append(np.mean(bootstrap_sample))

    mean_estimate = np.mean(mean_estimates)
    return mean_estimate

# Define the sample_data, num_samples, and sample_size
sample_data = np.array([165, 170, 175, 180, 185])  # Example data
num_samples = 1000  # Number of bootstrap samples to generate
sample_size = 100  # Size of each bootstrap sample

# Call the function
mean_estimate = bootstrap_estimate_height(sample_data, num_samples, sample_size)
print(mean_estimate)
#raise NotImplementedError()
bootstrap_estimate_height(sample_data, num_samples, sample_size)


```

Output:
```
174.9909
175.01364999999998
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell

# Sample data (heights of students)
sample_data = np.array([160, 165, 170, 175, 180, 185, 190, 160, 165, 170, 175, 180, 185, 190])

# Parameters
num_samples = 1000
sample_size = len(sample_data)

tolerance = 0.05
assert np.allclose(bootstrap_estimate_height(sample_data, num_samples, sample_size), 175.02, rtol=tolerance)
```

---
## Question 2:

#### Majority Voting
Given a list = `[1,2,3,2,2,1,3,3,3,2,2,1]`, consider 12 judges sitting and deciding among 3 contenders. Write a function to find the winner-based on a majority voting scheme.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
def my_majorityvote(a):
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    """
    Find the winner based on a majority voting scheme.

    Parameters:
    - a (list): List of votes.

    Returns:
    - winner: The winner based on the majority of votes.
    """
    from collections import Counter

    # Count the occurrences of each contender
    vote_counts = Counter(a)

    # Find the contender with the maximum number of votes
    winner = max(vote_counts, key=vote_counts.get)

    return winner
#raise NotImplementedError()

```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
# Test case 1
a = np.array([1,2,3,2,2,1,3,3,3,2,2,1])
assert my_majorityvote(a) == 2

# Test case 2
b = np.array([1,2,2,2,3,3,3,3,3,1])
assert my_majorityvote(b) == 3
```

---
## Question 3:
The `Gini index` is a measure of impurity or inequality used in `decision tree` algorithms. It quantifies the likelihood of misclassifying a randomly chosen element if it were randomly classified according to the distribution of labels in the node. The formula is given as follows:


<img src="https://www.learnbymarketing.com/wp-content/uploads/2016/02/gini-index-formula.png" alt="drawing" style="width:200px;"/>

Write a function `calculate_gini_index` calculate `Gini index`. Labels is a list of class labels. Your function `calculate_gini_index` should calculates the `Gini index` based on the frequency of each class label in the dataset. Finally, it returns the Gini index value as an output.


<!-- grade_id: my_GiniIdex, points: 0, grade: false, solution: true -->
```python
def calculate_gini_index(labels):
    """
    Calculate the Gini index for a set of class labels.

    Parameters:
    - labels (list or numpy.ndarray): List of class labels.

    Returns:
    - gini_index (float): Gini index.
    """
    # YOUR CODE HERE
from collections import Counter

def calculate_gini_index(labels):
    """
    Calculate the Gini index for a set of class labels.

    Parameters:
    - labels (list or numpy.ndarray): List of class labels.

    Returns:
    - gini_index (float): Gini index.
    """
    # Count the occurrences of each class label
    label_counts = Counter(labels)

    # Calculate the total number of samples
    total_samples = len(labels)

    # Calculate the Gini index
    gini_index = 1.0
    for label in label_counts:
        label_prob = label_counts[label] / total_samples
        gini_index -= label_prob ** 2

    return gini_index

# Define your list of class labels
labels = [1, 1, 2, 2, 2, 3]

# Call the function with your list of class labels
gini_index = calculate_gini_index(labels)
print("Gini index:", gini_index)
#raise NotImplementedError()


```

Output:
```
Gini index: 0.611111111111111
```

<!-- grade_id: my_GiniIdex_solution, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell

# Test case 1
labels = [1, 0, 1, 1, 1, 1, 0, 1]
gini_index = calculate_gini_index(labels)
assert gini_index == 0.375


# Test case 2
labels = [1, 0, 0, 0, 1, 1, 0, 1]
gini_index = calculate_gini_index(labels)
assert gini_index == 0.5

```