<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/problem2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook
### Autumn Term 2023-24
#### Prepared by Dr Haider Raza

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a Python function that takes a list of numbers and returns a list of their cosine values.

By solving this question correctly you will get 1 point.

<!-- grade_id: cosine_ans, points: 0, grade: false, solution: true -->
```python
import math
def calculate_cosine_of_list(numbers):
    """Write a Python function that takes a list of numbers and returns a list of their cosine values """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
numbers = [0, -2, 44, -61, -30, 44]
cosine_result = calculate_cosine_of_list(numbers)
print(cosine_result)
```

Output:
```
<ipython-input-2-3ddfb77a5c8f> in <cell line: 2>()
      1 numbers = [0, -2, 44, -61, -30, 44]
----> 2 cosine_result = calculate_cosine_of_list(numbers)
      3 print(cosine_result)

<ipython-input-1-d51b518d87eb> in calculate_cosine_of_list(numbers)
      3     """Write a Python function that takes a list of numbers and returns a list of their cosine values """
      4     # YOUR CODE HERE
----> 5     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cosine_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""

# Test case 1
my = [0, math.pi / 2, math.pi, 100, 200, 230]
assert calculate_cosine_of_list(my) == [1.0, 0.0, -1.0, 0.86, 0.49, -0.79]

# Test case 2
my = [0, 2, 4, 6, -30, 400]
assert calculate_cosine_of_list(my) == [1.0, -0.42, -0.65, 0.96, 0.15, -0.53]
```

---

## Question 2:

Given an array of integers, find two numbers such that they add up to a specific target sum. Return the indices of the two numbers. For example, given the array `[2, 7, 11, 15]` and the target `9`, the function should return `[0, 1]` since `2 + 7 = 9`.

By solving this question correctly you will get 1 point.

<!-- grade_id: twosum_ans, points: 0, grade: false, solution: true -->
```python
def two_sum(nums, target):
    """Given an array of integers, find two numbers such that they add
    up to a specific target sum. Return the indices of the two numbers."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Example for students
nums = [2, 7, 11, 15]
target = 9
result = two_sum(nums, target)
print(result)
```

<!-- grade_id: twosum_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""

# Case 1
nums = [17, 12, 31, 6]
target = 37
assert two_sum(nums, target) == [2, 3]

# Case 2
nums = [1, 2, 4, 6, 30, 40]
target = 31
assert two_sum(nums, target) == [0, 4]
```

---
## Question 3:

Given an array containing all integers from `0` to `n` except for one missing number, find the missing number. For example, given `[3, 0, 1]`, the missing number is `2`.

By solving this question correctly you will get 1 point.

<!-- grade_id: missnum_ans, points: 0, grade: false, solution: true -->
```python

def find_missing_number(nums):
    """Given an array containing all integers from 0 to n except for one missing number,
    find the missing number. For example, given [3, 0, 1],
    the missing number is 2."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
nums = [3, 0, 1]
result = find_missing_number(nums)
print(result)
```

<!-- grade_id: missnum_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
# Case 1
nums = [3, 0, 1]
assert find_missing_number(nums) == 2

# Case 2
nums = [4, 5, 6, 2, 0, 1]
assert find_missing_number(nums) == 3

```

## Pandas Data Frame


```python
from sklearn import datasets
import pandas as pd
# Create a sample DataFrame
data = {
    'Name': ['Alice', 'Bob', 'Charlie', 'David'],
    'Age': [25, 30, 28, 35],
    'City': ['New York', 'Los Angeles', 'Chicago', 'Houston']
}
df = pd.DataFrame(data)
df
```

---
## Question 4:

Write a function to Find the age of the person named `Bob` and `Bob` is not in the Data Frame then return `0`.

By solving this question correctly you will get 2 points.

<!-- grade_id: findele_ans, points: 0, grade: false, solution: true -->
```python
def find_element(dataframe, name):
    """Write a function to Find the age of the person named "Bob" """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: findele_test, points: 2, grade: true, solution: false -->
```python
assert find_element(df, 'Bob') == 30
assert find_element(df, 'Charlie') == 28
assert find_element(df, 'CE880') == 0
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of `Age`. `Hint` Create a list and append value to it and return the output.

By solving this question correctly you will get 1 point.

```python

```

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of Age from the
    above dataframe df and round the values to 2 decimals points ."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: stats_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
x = stats(df)
y= [29.5, 4.2]
np.testing.assert_allclose(x, y, rtol=0.05, atol=0)
```