<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Siham Anaya"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

We have already loaded `iris` dataset for you. Now, write a function to calculate the `median` on the 1st column of the dataset.
`Hint`: To find the median, you first need to sort the values in your sample. But, before that we have sliced, the first column of `iris` data into a variable name `sample` and data type is a list. Then pass `sample` list as a parameter in you `my_median` function.

`Note`: Please don't use `np.median` function. Write your own custom function

<!-- grade_id: my_median_ans, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets
import numpy as np
import pandas as pd

iris = datasets.load_iris()
sample = list(iris.data[:,0])

def my_median(sample):
    """Function for calculating median"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_median_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
import math
assert math.isclose(my_median(sample), 5.8, rel_tol=0.05)
```

---

## Question 2:

Write a function to Caluate Interquartile Range (IQR) on the second column of `iris` dataset, which you have loaded in the question 1.
`hint`: We have created a new variable `sample2` and sliced the second column of the `iris` data. Please explore `np.percentile` function

<!-- grade_id: my_quartile_ans, points: 0, grade: false, solution: true -->
```python

sample2 = list(iris.data[:,1])

def my_quartile(sample2):
    """Caluate Interquartile Range of the given list."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_quartile_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct answers here"""
import math
assert math.isclose(my_quartile(sample2), 0.5, rel_tol=0.05)
```

---
## Question 3:

Create a function that, given a data frame `df`, returns the name feature with the highest standard deviation across all species. For instance, if the function is called `feature_with_highest_std(data):`," it should return `pl` as the feature with the highest standard deviation, where `pl` stands for `petal length (cm)`.

<!-- grade_id: ans_high_std, points: 0, grade: false, solution: true -->
```python
import math

from sklearn.datasets import load_iris
import pandas as pd
import numpy as np

data = load_iris()
df = pd.DataFrame(data=data.data, columns=data.feature_names)
df = df.rename(columns={"sepal length (cm)": "sl", "sepal width (cm)": "sw",
                        "petal length (cm)":"pl", "petal width (cm)": "pw" })
df.head()

def feature_with_highest_std(data):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: my_std_test, points: 2, grade: true, solution: false -->
```python
"""Test you solution"""
assert feature_with_highest_std(df) == 'pl'
```

---
## Question 4:

You have been given the `Iris` dataset from the `sklearn` library. Write a Python function `custom_mean(dataset, feature_name)` that takes an `Iris` dataset and a `feature name` as input and returns the custom mean of that feature. The custom mean is calculated by: Removing the top 10% of the highest values and the bottom 10% of the lowest values for that feature. Calculating the mean of the remaining values and round it to two decimal points.

<!-- grade_id: my_mean_ans, points: 0, grade: false, solution: true -->
```python
def custom_mean(data, feature_name):
    """Create Python function custom_mean using Iris dataset,
    input feature name, and calculate mean after removing extremes."""
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: my_mean_test, points: 2, grade: true, solution: false -->
```python
"""Test you solution"""
assert custom_mean(df, 'sl') == 5.79
assert custom_mean(df, 'sw') == 3.03
assert custom_mean(df, 'pl') == 3.72
assert custom_mean(df, 'pw') == 1.12
```

```python

```