<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week1/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a function for Pythagorean theorem, or Pythagoras' theorem. Pythagorean is a fundamental relation in Euclidean geometry among the three sides of a right triangle. It states that the area of the square whose side is the hypotenuse (the side opposite the right angle) is equal to the sum of the areas of the squares on the other two sides. This theorem can be written as an equation relating the lengths of the legs a, b and the hypotenuse c, often called the Pythagorean equation:

$a^2 + b^2 = c^2$

Please write a function (code)to calculate the length of an unknown side of a right triangle, given two sides are known.

<!-- grade_id: myPythagorean_ans, points: 0, grade: false, solution: true -->
```python
import math
def myPythagorean(a,b):
    if a == 0 or b == 0:
        return 0
    else:
        c = math.sqrt(a**2 + b**2)  # applying the Pythagorean theorem
        return c
    raise NotImplementedError()
```

```python
myPythagorean(10,14)
```

Output:
```
17.204650534085253
```

<!-- grade_id: myPythagorean_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert myPythagorean(3,2) == 3.605551275463989
assert myPythagorean(7,3) == 7.615773105863909
assert myPythagorean(0,2) == 0
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />
     
Please write a function (code)to calculate the area of trapezoid.

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def calcul_area(a, b, h):

    area = 0.5 * (a + b) * h
    return area
base1 = 5
base2 = 3
height = 6
trapezoid_area = calcul_area(base1, base2, height)
print("The area of the trapezoid is:", trapezoid_area)
```

Output:
```
The area of the trapezoid is: 24.0
```

```python
calcul_area (5, 3, 6)
```

Output:
```
24.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert calcul_area(10, 10, 20) == 200.0
assert calcul_area(5, 10, 20) == 150.0
assert calcul_area(5, 3, 5) == 20.0
```

---
## Question 3:

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string2 (str1):
  m=len(str1)-1
  x= ""
  for i in range(m,-1,-1) :
    x=x+str1[i]
  return x



def reverse_string(str1):
    return str1[::-1]

```

```python
reverse_string ('ChatGPT')
reverse_string2 ('ChatGPT')

```

Output:
```
'TPGtahC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Humanitarian') == 'nairatinamuH'
assert reverse_string('ChatGPT') == 'TPGtahC'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string2('Humanitarian') == 'nairatinamuH'
assert reverse_string2('ChatGPT') == 'TPGtahC'
assert reverse_string2('Data Science') == 'ecneicS ataD'
assert reverse_string2('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the even numbers from a given list and in ascending order.

Sample List : `[1, 2, 3, 4, 5, 6, 7, 8, 9]` Expected Result : `[2, 4, 6, 8]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def even_num(lst_1):
    even_numbers = list(set([num for num in lst_1 if num % 2 == 0]))  # Remove duplicates and filter even numbers
    even_numbers.sort()  # Sort the list in ascending order
    return even_numbers



```

```python
numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9]
numbers2 = []
for X in numbers:
  if X%2==0:
    numbers2.append (X)
print (numbers2)

```

Output:
```
[2, 4, 6, 8]
```

```python
def even_num2 (numbers):
  numbers2 = []
  for X in numbers:
    if X%2==0 and X not in numbers2:
      numbers2.append (X)
    numbers2.sort()
  return numbers2
      #havent worked with sets but lists
```

```python
even_num([2, 3, 14, 17, 8, 19, 11, 14])
print(even_num2([1, 2, 3, 4, 6, 7, 8, 9]))
```

Output:
```
[2, 4, 6, 8]
```

<!-- grade_id: odd_num_test, points: 1, grade: true, solution: false -->
```python
assert even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
assert even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
assert even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [2, 8, 14]
```

```python
assert even_num2([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
assert even_num2([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
assert even_num2([2, 3, 14, 17, 8, 19, 11, 14]) == [2, 8, 14]
```

## Question 5:
Write a Python function to create and print a list where the values are square of numbers between `1` and `n`, where `n` is user defined number: Example `sqr_values(4) == [1, 4, 9, 16]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):
    """Write a Python program to print the square of numbers between 1 and `n`"""

    # YOUR CODE HERE
    raise NotImplementedError()

def sqr_values(n):
    """
    Create and print a list of square values of numbers from 1 to n.

    Parameters:
    n (int): Upper limit for the range of numbers.

    Returns:
    list: List of square values.
    """
    result = [i*i for i in range(1, n+1)]
    return result



```

```python
def sqr_values2 (number):
  b= []
  for i in range (1,number+1):
    b.append (i**2)
  return (b)

```

Output:
```
give a number9
[1, 4, 9, 16, 25, 36, 49, 64, 81]
```

```python
sqr_values(4)
```

Output:
```
[1, 4, 9, 16]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9, 16]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64, 81]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

```python
assert sqr_values2(4) == [1, 4, 9, 16]
assert sqr_values2(8) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values2(9) == [1, 4, 9, 16, 25, 36, 49, 64, 81]
assert sqr_values2(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

## Question 6:

Write a Python program to find all the duplicate characters which are similar to each other and join them.


<!-- grade_id: comb_ans, points: 0, grade: false, solution: true -->
```python
def findDupChar(input_string):
    """
    Find and join all the duplicate characters that are similar to each other in a string.

    Parameters:
    input_string (str): Input string to process.

    Returns:
    str: Resulting string after joining similar duplicates.
    """
    input_string = input_string.lower()
    seen = set()
    result = []
    for char in input_string:
        if char in seen and char not in result:
            result.append(char)
        seen.add(char)
    return ' '.join((result))



```

```python
def find_input(str1):
    str1 = str1.lower()
    seen = set()
    for c in str1:
        seen.add(c)
    return seen



# Python code to count the number of occurrences
def countX(lst, x):
    count = 0
    for ele in lst:
        if (ele == x):
            count = count + 1
    return count


def findDupChar2(input_string):
    b=input_string
    result= set()
    for c in b:
        if countX(b,c)>1:
          result.add(c)
    return ' '.join((result))
```

```python
findDupChar('Essex University')
findDupChar2('Essex University')
```

Output:
```
'e s i'
```

<!-- grade_id: comb_test, points: 1, grade: true, solution: false -->
```python
assert findDupChar('Happyy') == 'p y'
assert findDupChar('Essex University') == 's e i'
```

```python
assert findDupChar2('Happyy') == 'p y'
assert findDupChar2('Essex University') == 's e i'

# WHY IS IT NOT IN ORDER ??????
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-140-dafed88fff17> in <cell line: 1>()
----> 1 assert findDupChar2('Happyy') == 'p y'
      2 assert findDupChar2('Essex University') == 's e i'

AssertionError:
```