<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week9/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-9 : Problem notebook


---
## Question 1:

#### Balancing Imbalance Dataset
An imbalanced dataset is a dataset where the distribution of `classes (or labels)` is skewed, meaning that one class has significantly more instances than the other class(es). This can lead to issues when training machine learning models, as they may become biased towards the majority class and perform poorly on predicting the minority class.

To balance an imbalanced dataset using a basic approach, one common technique is random undersampling, where instances from the majority class are randomly removed until the dataset is balanced. Write a Python function to perform random undersampling:

<!-- grade_id: my_model_ans, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np

# Read the CSV file into a DataFrame
data = pd.read_csv('ParisHousingClass.csv')

# Display the first few rows of the DataFrame
print(data.head())


def balance_data(X, y):
    """
    Balance the dataset by randomly undersampling the majority class.

    Parameters:
    - X (numpy.ndarray): Feature data.
    - y (numpy.ndarray): Target labels.

    Returns:
    - X_balanced (numpy.ndarray): Balanced feature data.
    - y_balanced (numpy.ndarray): Balanced target labels.
    """
    # YOUR CODE HERE

    # Find the indices of each class
    indices_class_0 = np.where(y == 0)[0]
    indices_class_1 = np.where(y == 1)[0]

    # Determine the minority and majority class based on the class indices' counts.
    if len(indices_class_1) < len(indices_class_0):
        minority_indices = indices_class_1
        majority_indices = indices_class_0
    else:
        minority_indices = indices_class_0
        majority_indices = indices_class_1

    # Set a random seed for reproducibility.
    np.random.seed(42)

    # Randomly sample from the majority class to match the minority class size.
    downsampled_majority_indices = np.random.choice(majority_indices, size=len(minority_indices), replace=False)

    # Combine the downsampled majority class indices with the minority class indices.
    balanced_indices = np.concatenate((minority_indices, downsampled_majority_indices))

    # Subset the original X and y based on the balanced indices.
    X_balanced = X[balanced_indices]
    y_balanced = y[balanced_indices]

    return X_balanced, y_balanced



#raise NotImplementedError()



```

Output:
```
squareMeters  numberOfRooms  hasYard  hasPool  floors  cityCode  \
0         75523              3        0        1      63      9373   
1         80771             39        1        1      98     39381   
2         55712             58        0        1      19     34457   
3         32316             47        0        0       6     27939   
4         70429             19        1        1      90     38045   

   cityPartRange  numPrevOwners  made  isNewBuilt  hasStormProtector  \
0              3              8  2005           0                  1   
1              8              6  2015           1                  0   
2              6              8  2021           0                  0   
3             10              4  2012           0                  1   
4              3              7  1990           1                  0   

   basement  attic  garage  hasStorageRoom  hasGuestRoom      price category  
0      4313   9005     956               0             7  7559081.5   
```

<!-- grade_id: balance_test1, points: 4, grade: true, solution: false -->
```python
# Passing this test case will give you 4 marks
# Test case with imbalanced dataset
X_test = np.array([[1, 2], [2, 3], [3, 4], [4, 5], [5, 6], [6, 7]])
y_test = np.array([0, 0, 1, 1, 1, 1])

X_balanced_test, y_balanced_test = balance_data(X_test, y_test)
print(y_balanced_test)
# Check if the balanced dataset has equal instances of each class
unique_classes, class_counts = np.unique(y_balanced_test, return_counts=True)
assert np.all(class_counts == class_counts[0]), "Balanced dataset should have equal instances of each class"
```

Output:
```
[0 0 1 1]
```

<!-- grade_id: balance_test2, points: 5, grade: true, solution: false -->
```python
# Passing this test case will give you 5 marks
# Test case with imbalanced dataset


X_test = np.array([[1, 2], [2, 3], [3, 4], [4, 5], [5, 6], [6, 7]])
y_test = np.array([0, 0, 0, 0, 0, 1])

X_balanced_test, y_balanced_test = balance_data(X_test, y_test)
print(y_balanced_test)
# Check if the balanced dataset has equal instances of each class
unique_classes, class_counts = np.unique(y_balanced_test, return_counts=True)
assert np.all(class_counts == class_counts[0]), "Balanced dataset should have equal instances of each class"
```

Output:
```
[1 0]
```