<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week9/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-9 : Problem notebook


---
## Question 1:

#### Balancing Imbalance Dataset
An imbalanced dataset is a dataset where the distribution of `classes (or labels)` is skewed, meaning that one class has significantly more instances than the other class(es). This can lead to issues when training machine learning models, as they may become biased towards the majority class and perform poorly on predicting the minority class.

To balance an imbalanced dataset using a basic approach, one common technique is random undersampling, where instances from the majority class are randomly removed until the dataset is balanced. Write a Python function to perform random undersampling:

<!-- grade_id: my_model_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def balance_data(X, y):
    """
    Balance the dataset by randomly undersampling the majority class.

    Parameters:
    - X (numpy.ndarray): Feature data.
    - y (numpy.ndarray): Target labels.

    Returns:
    - X_balanced (numpy.ndarray): Balanced feature data.
    - y_balanced (numpy.ndarray): Balanced target labels.
    """
    # YOUR CODE HERE

def balance_data(X, y):
    """
    Balance an imbalanced dataset using random undersampling.

    Parameters:
    - X (numpy.ndarray): Feature matrix.
    - y (numpy.ndarray): Array of class labels.

    Returns:
    - X_resampled (numpy.ndarray): Resampled feature matrix.
    - y_resampled (numpy.ndarray): Resampled array of class labels.
    """
    # Find the majority class label
    majority_class_label = np.argmax(np.bincount(y))

    # Find indices of majority class instances
    majority_indices = np.where(y == majority_class_label)[0]

    # Find indices of minority class instances
    minority_indices = np.where(y != majority_class_label)[0]

    # Randomly select instances from majority class
    num_minority_instances = len(minority_indices)
    sampled_majority_indices = np.random.choice(majority_indices, size=num_minority_instances, replace=False)

    # Concatenate minority class instances with sampled majority class instances
    resampled_indices = np.concatenate([minority_indices, sampled_majority_indices])

    # Shuffle resampled indices
    np.random.shuffle(resampled_indices)

    # Return resampled dataset
    X_resampled = X[resampled_indices]
    y_resampled = y[resampled_indices]

    return X_resampled, y_resampled
    raise NotImplementedError()

```

<!-- grade_id: balance_test1, points: 4, grade: true, solution: false -->
```python
# Passing this test case will give you 4 marks
# Test case with imbalanced dataset
X_test = np.array([[1, 2], [2, 3], [3, 4], [4, 5], [5, 6], [6, 7]])
y_test = np.array([0, 0, 1, 1, 1, 1])

X_balanced_test, y_balanced_test = balance_data(X_test, y_test)
print(y_balanced_test)
# Check if the balanced dataset has equal instances of each class
unique_classes, class_counts = np.unique(y_balanced_test, return_counts=True)
assert np.all(class_counts == class_counts[0]), "Balanced dataset should have equal instances of each class"
```

Output:
```
[0 0 1 1]
```

<!-- grade_id: balance_test2, points: 5, grade: true, solution: false -->
```python
# Passing this test case will give you 5 marks
# Test case with imbalanced dataset


X_test = np.array([[1, 2], [2, 3], [3, 4], [4, 5], [5, 6], [6, 7]])
y_test = np.array([0, 0, 0, 0, 0, 1])

X_balanced_test, y_balanced_test = balance_data(X_test, y_test)
print(y_balanced_test)
# Check if the balanced dataset has equal instances of each class
unique_classes, class_counts = np.unique(y_balanced_test, return_counts=True)
assert np.all(class_counts == class_counts[0]), "Balanced dataset should have equal instances of each class"
```

Output:
```
[0 1]
```

```python

```