<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week4/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "SIHAM ANAYA"
COLLABORATORS = "2316969"
```

---

---
# Welcome to CE880
### This is your week-4 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

### Task: Complete all the three questions given as follows:

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
# !pip install -U scikit-learn # Run if required
```

```python
import pandas as pd
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn.cluster import KMeans
from sklearn.datasets import make_circles, make_moons
from sklearn.metrics import adjusted_rand_score
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score, adjusted_rand_score
from sklearn.decomposition import PCA
scaler = StandardScaler()
```

---
## Question 1:
We have already loaded `moon` data for you using `make_moons` function.

Write a code to perform `k-means` clustering on `scaled_features` variable and return `silhouette_score` as an output

<!-- grade_id: myScore_kmean_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.cluster import KMeans
from sklearn.metrics import silhouette_score
from sklearn.datasets import make_moons
from sklearn.preprocessing import StandardScaler

# Let's make two moons, where n=250 is the number of samples
features, true_labels = make_moons(n_samples=250, noise=0.05, random_state=42)
# Initialize a StandardScaler
scaler = StandardScaler()
# fit_transform() method is basically the combination of fit method and transform method,
# it is equivalent to fit().transform(). This method performs fit and transform on the
# input data at a single time and converts the data points.
# If we use fit and transform separate when we need both then it
# will decrease the efficiency of the model so we use fit_transform()
# which will do both the work.
scaled_features = scaler.fit_transform(features)
myscore = []

def myScore_kmean(scaled_features):
    """Write code to perform k-means clustering and print silhouette_score"""
    # YOUR CODE HERE
    for k in range(2, 11):
        kmeans = KMeans(n_clusters=k, random_state=42)
        labels = kmeans.fit_predict(scaled_features)
        score = silhouette_score(scaled_features, labels)
        myscore.append(score)

    return max(myscore)

    raise NotImplementedError()





```

<!-- grade_id: myScore_kmean_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(myScore_kmean(scaled_features), 0.5, rel_tol=0.10)
```

Output:
```
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
/usr/local/lib/python3.10/dist-packages/sklearn/cluster/_kmeans.py:8
```

---
## Question 2:

Write a code to perform `DBSCAN` clustering algorithm on `scaled_features` and return `silhouette_score` as an output. Please set the `eps` = 0.2 and `min_samples` = 5

<!-- grade_id: myScore_DBSCAN_ans, points: 0, grade: false, solution: true -->
```python

from sklearn.datasets import make_moons
from sklearn.cluster import DBSCAN
from sklearn.metrics import silhouette_score
from sklearn.preprocessing import StandardScaler

# Let's make two moons, where n is the number of samples
features, true_labels = make_moons(n_samples=250, noise=.1)

def myScore_DBSCAN(scaled_features):
    """Write code to perform DBSCAN algorithm and print silhouette_score"""
    # YOUR CODE HERE
    eps = 0.2
    min_samples = 5

    dbscan = DBSCAN(eps=eps, min_samples=min_samples)
    labels = dbscan.fit_predict(scaled_features)

    if len(set(labels)) <= 1:  # If all points are in the same cluster or classified as noise
        return -1  # Return -1 for silhouette score

    score = silhouette_score(scaled_features, labels)

    return score
    raise NotImplementedError()




```

<!-- grade_id: myScore_DBSCAN_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert math.isclose(myScore_DBSCAN(scaled_features), 0.2, rel_tol=0.10)
```

## Question 3:

Write a Python function `compute_covariance_matrix(X)` that calculates the covariance matrix of the input data.

Input: `X` (A numpy array representing the input data matrix of shape `(n_samples, n_features)`.

Output: A numpy array representing the covariance matrix of shape `(n_features, n_features)`.

Help: [Covariance](https://en.wikipedia.org/wiki/Covariance)

<!-- grade_id: mycov_answ, points: 0, grade: false, solution: true -->
```python
import numpy as np
def compute_covariance_matrix(X):
    """
    Compute the covariance matrix of the input data.

    Parameters:
    X (numpy.ndarray): Input data matrix of shape (n_samples, n_features).

    Returns:
    numpy.ndarray: Covariance matrix of shape (n_features, n_features).
    """
    # YOUR CODE HERE
    # Compute the covariance matrix
    covariance_matrix = np.cov(X, rowvar=False)

    return covariance_matrix

    raise NotImplementedError()

```

<!-- grade_id: cell-mycov_test, points: 2, grade: true, solution: false -->
```python
import numpy as np
array2 = [[ 0.08177146,  0.00473089, -0.0073587 ],
          [ 0.00473089,  0.09623581, -0.0013255 ],
          [-0.0073587,  -0.0013255,   0.08300967]]
np.random.seed(42)
X = np.random.rand(100, 3)  # 100 samples, 3 features
cov_matrix = compute_covariance_matrix(X)

assert np.allclose(cov_matrix, array2, rtol=1e-05)
```