<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week3/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a Python function to find the mode(s) in a `list` of integers. The mode is the number that appears most frequently. However, the challenge here is that the list may contain multiple `modes`, and your task is to efficiently find and return all modes.


<!-- grade_id: my_mode, points: 0, grade: false, solution: true -->
```python

def find_modes(numbers):
    """Python function to find the mode(s) in a list of integers.
    Python function to find the mode(s) in a list of integers.

    Parameters:
        numbers (list of int): The input list of integers.

    Returns:
        list: A list containing the mode(s) in the input list.
    """
    from collections import Counter  # Importing Counter within the function

    # Count the occurrences of each number
    num_counts = Counter(numbers)

    # Find the maximum count
    max_count = max(num_counts.values())

    # Find all numbers with the maximum count (modes)
    modes = [num for num, count in num_counts.items() if count == max_count]

    return modes
    raise NotImplementedError()


```

```python
numbers = [3, 5, 3, 3, 1, 5, 2]
find_modes(numbers)
```

Output:
```
[3]
```

<!-- grade_id: my_mode_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
numbers = [2, 3, 5, 2, 8, 7, 2, 5, 4, 6, 5, 2]
assert find_modes(numbers) == [2]

numbers = [2, 3, 5, 2, 3, 3, 2, 3, 4, 6, 5, 2]
assert find_modes(numbers) == [2, 3]
```

---

## Question 2:

Write a Python function to calculate the `Interquartile Range (IQR)` for a given list of numerical values.

<!-- grade_id: my_quartile_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def calculate_iqr(data):
    """Write a Python function to calculate the `Interquartile Range (IQR)`

    Python function to calculate the Interquartile Range (IQR) for a given list of numerical values.

    Parameters:
        data (list of float or int): The input list of numerical values.

    Returns:
        float: The Interquartile Range (IQR) of the input data.
    """
    # Calculate the first quartile (Q1)
    q1 = np.percentile(data, 25)

    # Calculate the third quartile (Q3)
    q3 = np.percentile(data, 75)

    # Calculate the Interquartile Range (IQR)
    iqr = q3 - q1

    return iqr

```

```python
data = [12, 15, 20, 25, 30, 15, 10, 15, 16, 15]
calculate_iqr(data)
```

Output:
```
4.0
```

<!-- grade_id: my_quartile_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct answers here"""
data = [12, 15, 20, 25, 30, 35, 40, 45, 50, 55]
assert calculate_iqr(data) == 25

data = [42, 55, 40, 35, 50, 45, 30, 45, 30, 55]
assert  calculate_iqr(data) == 15
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-14-693e7deb0247> in <cell line: 3>()
      1 """Check that correct answers here"""
      2 data = [12, 15, 20, 25, 30, 35, 40, 45, 50, 55]
----> 3 assert calculate_iqr(data) == 25
      4 
      5 data = [42, 55, 40, 35, 50, 45, 30, 45, 30, 55]

AssertionError:
```

---
## Question 3:

Create a function that, given a data frame `df`, returns the name feature with the highest standard deviation across all species. For instance, if the function is called `feature_with_highest_std(data):`," it should return `pl` as the feature with the highest standard deviation, where `pl` stands for `petal length (cm)`.

<!-- grade_id: ans_high_std, points: 0, grade: false, solution: true -->
```python
import math
from sklearn.datasets import load_iris
import pandas as pd
import numpy as np

data = load_iris()
df = pd.DataFrame(data=data.data, columns=data.feature_names)
df = df.rename(columns={"sepal length (cm)": "sl", "sepal width (cm)": "sw",
                        "petal length (cm)":"pl", "petal width (cm)": "pw" })
df.head()

def feature_with_highest_std(data):
    """
    Given a data frame df, returns the name feature with the highest standard deviation across all species.

    Parameters:
        data (DataFrame): The input DataFrame.

    Returns:
        str: The name of the feature with the highest standard deviation.
    """
    # Calculate standard deviation for each feature
    std_devs = data.std()

    # Find the feature with the highest standard deviation
    highest_std_feature = std_devs.idxmax()

    return highest_std_feature

    raise NotImplementedError()



```

<!-- grade_id: my_std_test, points: 2, grade: true, solution: false -->
```python
"""Test you solution"""
assert feature_with_highest_std(df) == 'pl'
```