<a href="https://colab.research.google.com/github/sihamanaya/CE880_Lab_coursework/blob/main/week5/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-5 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

---
## Question 1:
Write a function to compute the `softmax` function for an input array `z`, representing the logits of a classification problem.

**Parameters**:
 `z` (`numpy.ndarray`): Input array of shape `(n_samples, n_classes)`.
    
**Returns**:
softmax_output `(numpy.ndarray)`: Output array after applying softmax along axis 1.



<img src="https://cdn.sanity.io/images/vr8gru94/production/582a6c51701bb584c1cdd6662cc376b9cadb7160-2048x1152.png" alt="drawing" style="width:400px;"/>

<!-- grade_id: my-softmax-ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def softmax(z):
    """
    Compute the softmax function for input array z.

    Parameters:
    - z (numpy.ndarray): Input array of shape (n_samples, n_classes).

    Returns:
    - softmax_output (numpy.ndarray): Output array after applying softmax along axis 1.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: my_regression_test, points: 3, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
z = np.array([[1.0, 2.0, 3.0],
              [2.0, 3.0, 4.0]])

z_ans = np.array([[0.09, 0.24, 0.67],
                  [0.09, 0.24, 0.67]])

assert np.allclose(z_ans, softmax(z), rtol=1e-05)
```

---
## Question 2:

Write Python function `knn_predict` implementing the k-nearest neighbors (KNN) classification algorithm: Psudo code is given below for your help:
For each test sample:
* a. Calculate the distance between the test sample and each training sample.
* b. Sort the distances in ascending order and select the k nearest neighbors.
* c. Determine the majority class among the k nearest neighbors.
* d. Assign the test sample to the majority class.

<!-- grade_id: my_RMregression_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np

def euclidean_distance(x1, x2):
    """
    Compute the Euclidean distance between two points.

    Parameters:
    - x1 (numpy.ndarray): First point.
    - x2 (numpy.ndarray): Second point.

    Returns:
    - distance (float): Euclidean distance between the points.
    """
    return np.sqrt(np.sum((x1 - x2)**2))

def knn_predict(X_train, y_train, X_test, k=3):
    """
    Predict the class labels for test samples using k-nearest neighbors algorithm.

    Parameters:
    - X_train (numpy.ndarray): Training features.
    - y_train (numpy.ndarray): Training labels.
    - X_test (numpy.ndarray): Test features.
    - k (int): Number of nearest neighbors to consider. Default is 3.

    Returns:
    - y_pred (list): Predicted labels for test samples.

    Psudo Code given as follows:
    For each test sample:
        a. Calculate the distance between the test sample and each training sample.
        b. Sort the distances in ascending order and select the k nearest neighbors.
        c. Determine the majority class among the k nearest neighbors.
        d. Assign the test sample to the majority class.

    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: my_KNN_test, points: 3, grade: true, solution: false -->
```python
# Check you solution by running this cell
X_train = np.array([[1, 2], [2, 3], [3, 4], [4, 5], [5, 6]])
y_train = [0, 0, 0, 1, 1]
X_test = np.array([[3, 2], [5, 5]])
y_test =[0, 1]

assert knn_predict(X_train, y_train, X_test, k=3) == y_test
```